<a href="https://colab.research.google.com/github/MedhatElhawy/solvingNotbooks/blob/main/Solution__week1_python_fundamentals.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---

### Now It's Your Turn! 🎯

Let's practice what you learned:

<div dir="rtl">

### دورك! 🎯


### Exercise 1.1: Create Variables


<!-- grade_id: ex1_1, points: 0, grade: false, solution: true -->
```python
# Create these variables:
# - weight: a float equal to 0.5
# - bias: a float equal to 0.0
# - num_features: an integer equal to 10
# - dataset_name: a string equal to "mnist"

# YOUR CODE HERE
weight = 0.5 # indentation
bias = 0.0
num_features = 10
dataset_name = "mnist"


print(f"Weight: {weight}")
print(f"Bias: {bias}")
print(f"Number of Features: {num_features}")
print(f"Dataset: {dataset_name}")
```

Output:
```
Weight: 0.5
Bias: 0.0
Number of Features: 10
Dataset: mnist
```

<!-- grade_id: ex1_1_test, points: 1, grade: true, solution: false -->
```python
assert weight == 0.5 and isinstance(weight, float)
assert bias == 0.0 and isinstance(bias, float)
assert num_features == 10 and isinstance(num_features, int)
assert dataset_name == "mnist" and isinstance(dataset_name, str)
print("✓ Exercise 1.1 passed!")
```

Output:
```
✓ Exercise 1.1 passed!
```

### Exercise 1.2: Type Conversion


<!-- grade_id: ex1_2, points: 0, grade: false, solution: true -->
```python
# Given these strings:
epochs_str = "500"
lr_str = "0.001"

# Convert them to numbers:
# - epochs_int: convert epochs_str to integer
# - lr_float: convert lr_str to float

# YOUR CODE HERE
epochs_int = int(epochs_str)
lr_float = float(lr_str)

print(f"Epochs: {epochs_int} (type: {type(epochs_int).__name__})")
print(f"Learning Rate: {lr_float} (type: {type(lr_float).__name__})")
```

Output:
```
Epochs: 500 (type: int)
Learning Rate: 0.001 (type: float)
```

<!-- grade_id: ex1_2_test, points: 1, grade: true, solution: false -->
```python
assert epochs_int == 500 and isinstance(epochs_int, int)
assert lr_float == 0.001 and isinstance(lr_float, float)
print("✓ Exercise 1.2 passed!")
```

Output:
```
✓ Exercise 1.2 passed!
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 2.1: List Basics


<!-- grade_id: ex2_1, points: 0, grade: false, solution: true -->
```python
# Create a list called 'predictions' with these values: 0, 1, 1, 0, 1, 0, 1, 1
# Then extract:
# - first_pred: the first element
# - last_pred: the last element (use negative indexing!)
# - third_pred: the third element

# YOUR CODE HERE
predictions = [0,1,1,0,1,0,1,1]
first_pred = predictions[0]
last_pred = predictions[-1]
third_pred = predictions[len(predictions) - 1]

print(f"Predictions: {predictions}")
print(f"First: {first_pred}, Third: {third_pred}, Last: {last_pred}")
```

Output:
```
Predictions: [0, 1, 1, 0, 1, 0, 1, 1]
First: 0, Third: 1, Last: 1
```

<!-- grade_id: ex2_1_test, points: 1, grade: true, solution: false -->
```python
assert predictions == [0, 1, 1, 0, 1, 0, 1, 1]
assert first_pred == 0 and third_pred == 1 and last_pred == 1
print("✓ Exercise 2.1 passed!")
```

Output:
```
✓ Exercise 2.1 passed!
```

### Exercise 2.2: List Slicing


<!-- grade_id: ex2_2, points: 0, grade: false, solution: true -->
```python
# Given:
losses = [0.8, 0.65, 0.52, 0.43, 0.38, 0.34, 0.31, 0.29, 0.27, 0.26]

# Extract:
# - first_three: first 3 values
# - last_three: last 3 values
# - middle: values from index 3 to 6 (inclusive of 3, exclusive of 7)

# YOUR CODE HERE
first_three = losses[:3]
last_three = losses[-3:]
middle= losses[3:7]

print(f"First three: {first_three}")
print(f"Last three: {last_three}")
print(f"Middle: {middle}")
```

Output:
```
First three: [0.8, 0.65, 0.52]
Last three: [0.29, 0.27, 0.26]
Middle: [0.43, 0.38, 0.34, 0.31]
```

<!-- grade_id: ex2_2_test, points: 1, grade: true, solution: false -->
```python
assert first_three == [0.8, 0.65, 0.52]
assert last_three == [0.29, 0.27, 0.26]
assert middle == [0.43, 0.38, 0.34, 0.31]
print("✓ Exercise 2.2 passed!")
```

Output:
```
✓ Exercise 2.2 passed!
```

### Exercise 2.3: List Operations


<!-- grade_id: ex2_3, points: 0, grade: false, solution: true -->
```python
# Start with:
feature_values = [1.2, 3.4, 2.1]

# Do these operations:
# 1. Append 4.5 to the list
# 2. Calculate length → store in 'num_features'
# 3. Calculate sum → store in 'total'

# YOUR CODE HERE
feature_values.append(4.5)
num_features = len(feature_values)
total = 0.0
for i in range(num_features):
  total+=feature_values[i]


print(f"Features: {feature_values}")
print(f"Count: {num_features}")
print(f"Total: {total}")
```

Output:
```
Features: [1.2, 3.4, 2.1, 4.5]
Count: 4
Total: 11.2
```

<!-- grade_id: ex2_3_test, points: 1, grade: true, solution: false -->
```python
assert feature_values == [1.2, 3.4, 2.1, 4.5]
assert num_features == 4 and abs(total - 11.2) < 0.01
print("✓ Exercise 2.3 passed!")
```

Output:
```
✓ Exercise 2.3 passed!
```

### Exercise 2.4: 2D Lists


<!-- grade_id: ex2_4, points: 0, grade: false, solution: true -->
```python
# Create a 2D list 'data' with 3 rows:
# Row 0: [170, 68]
# Row 1: [165, 62]
# Row 2: [180, 75]
#
# Then extract:
# - first_sample: entire first row
# - second_weight: weight from second row (index 1, then 1)

# YOUR CODE HERE
data = [
    [170, 68],
    [165, 62],
    [180, 75]
]
first_sample = data[0]
second_weight = data[1][1]

print(f"Data: {data}")
print(f"First sample: {first_sample}")
print(f"Second weight: {second_weight}")
```

Output:
```
Data: [[170, 68], [165, 62], [180, 75]]
First sample: [170, 68]
Second weight: 62
```

<!-- grade_id: ex2_4_test, points: 1, grade: true, solution: false -->
```python
assert data == [[170, 68], [165, 62], [180, 75]]
assert first_sample == [170, 68] and second_weight == 62
print("✓ Exercise 2.4 passed!")
```

Output:
```
✓ Exercise 2.4 passed!
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 3.1: Basic Loop


<!-- grade_id: ex3_1, points: 0, grade: false, solution: true -->
```python
predictions = [0, 1, 1, 0, 1]

# Use a loop to create strings like "Sample 0: Class 0"
# Store all strings in output_lines list

# YOUR CODE HERE
output_lines = []
for i in range(len(predictions)):
  output_lines.append(f"Sample {i}: Class {predictions[i]}")
```

<!-- grade_id: ex3_1_test, points: 1, grade: true, solution: false -->
```python
expected = ["Sample 0: Class 0", "Sample 1: Class 1", "Sample 2: Class 1", "Sample 3: Class 0", "Sample 4: Class 1"]
assert output_lines == expected
print("✓ Exercise 3.1 passed!")
```

Output:
```
✓ Exercise 3.1 passed!
```

### Exercise 3.2: Counting with Loops


<!-- grade_id: ex3_2, points: 0, grade: false, solution: true -->
```python
predictions = [0, 1, 1, 0, 1, 0, 1, 1]
true_labels = [0, 1, 0, 0, 1, 0, 1, 1]

# Count correct predictions
# Store in correct_count

# YOUR CODE HERE
correct_count = 0
for i in range(len(predictions)):
  if predictions[i] == true_labels[i]:
    correct_count += 1
print(f"Correct: {correct_count}/{len(predictions)}")
print(f"Accuracy: {correct_count/len(predictions)*100:.1f}%")
```

Output:
```
Correct: 7/8
Accuracy: 87.5%
```

<!-- grade_id: ex3_2_test, points: 1, grade: true, solution: false -->
```python
assert correct_count == 7
print("✓ Exercise 3.2 passed!")
```

Output:
```
✓ Exercise 3.2 passed!
```

### Exercise 3.3: Calculate Column Means


<!-- grade_id: ex3_3, points: 0, grade: false, solution: true -->
```python
# Data: 3 people with [height, weight]
import numpy as np
data = [
    [170, 68],
    [165, 62],
    [180, 75]
]
npArr = np.array(data)
# Calculate mean of each column
# mean_height: average of first column
# mean_weight: average of second column

# YOUR CODE HERE
mean_height = np.mean(npArr[:,0])
mean_weight = np.mean(npArr[:,1])

#mean_weight

print(f"Mean height: {mean_height:.2f} cm")
print(f"Mean weight: {mean_weight:.2f} kg")
```

Output:
```
Mean height: 171.67 cm
Mean weight: 68.33 kg
```

<!-- grade_id: ex3_3_test, points: 1, grade: true, solution: false -->
```python
assert abs(mean_height - 171.67) < 0.1 and abs(mean_weight - 68.33) < 0.1
print("✓ Exercise 3.3 passed!")
```

Output:
```
✓ Exercise 3.3 passed!
```

### Exercise 3.4: Nested Loop


<!-- grade_id: ex3_4, points: 0, grade: false, solution: true -->
```python
data = [
    [1, 2, 3],
    [4, 5, 6],
    [7, 8, 9]
]

# Sum ALL values using nested loops
# Store in total_sum

# YOUR CODE HERE
total_sum = 0
for i in range(len(data)):
  for j in range(len(data[i])):
    total_sum += data[i][j]
print(f"Total sum: {total_sum}")
```

Output:
```
Total sum: 45
```

<!-- grade_id: ex3_4_test, points: 1, grade: true, solution: false -->
```python
assert total_sum == 45
print("✓ Exercise 3.4 passed!")
```

Output:
```
✓ Exercise 3.4 passed!
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 4.1: Mean Function


<!-- grade_id: ex4_1, points: 0, grade: false, solution: true -->
```python
def calculate_mean(values):
    """
    Calculate mean (average) of a list.

    Args:
        values: List of numbers

    Returns:
        Mean as float
    """
    # YOUR CODE HERE
    mean = sum(values)/len(values) *(1.00)
    return mean


# Test
test = [10, 20, 30, 40, 50]
result = calculate_mean(test)
print(f"Mean of {test}: {result}")
```

Output:
```
Mean of [10, 20, 30, 40, 50]: 30.0
```

<!-- grade_id: ex4_1_test, points: 1, grade: true, solution: false -->
```python
assert abs(calculate_mean([1, 2, 3, 4, 5]) - 3.0) < 0.01
assert abs(calculate_mean([10, 20, 30]) - 20.0) < 0.01
print("✓ Exercise 4.1 passed!")
```

Output:
```
✓ Exercise 4.1 passed!
```

### Exercise 4.2: Distance Function


<!-- grade_id: ex4_2, points: 0, grade: false, solution: true -->
```python
def euclidean_distance(point1, point2):
    """
    Calculate Euclidean distance between two 2D points.

    Formula: sqrt((x1-x2)² + (y1-y2)²)
    """
    # YOUR CODE HERE
    x1, y1 = point1
    x2, y2 = point2

    dist = ((x1-x2)**2 + (y1-y2)**2)**0.5
    return dist

# Test
p1, p2 = [0, 0], [3, 4]
dist = euclidean_distance(p1, p2)
print(f"Distance: {dist}")
```

Output:
```
Distance: 5.0
```

<!-- grade_id: ex4_2_test, points: 1, grade: true, solution: false -->
```python
assert abs(euclidean_distance([0, 0], [3, 4]) - 5.0) < 0.01
assert abs(euclidean_distance([0, 0], [1, 1]) - 1.414) < 0.01
print("✓ Exercise 4.2 passed!")
```

Output:
```
✓ Exercise 4.2 passed!
```

### Exercise 4.3: Normalization Function


<!-- grade_id: ex4_3, points: 0, grade: false, solution: true -->
```python
def normalize_list(values):
    """
    Normalize to [0,1] using min-max scaling.

    Formula: (x - min) / (max - min)
    """

    # YOUR CODE HERE
    Min = min(values)
    Max = max(values)
    range = Max - Min

    normalized = []
    for val in values:
      normal_val = (val - Min) / range
      normalized.append(normal_val)
    return normalized

# Test
data = [10, 20, 30, 40, 50]
result = normalize_list(data)
print(f"Normalized: {[f'{x:.2f}' for x in result]}")
```

Output:
```
Normalized: ['0.00', '0.25', '0.50', '0.75', '1.00']
```

<!-- grade_id: ex4_3_test, points: 1, grade: true, solution: false -->
```python
r = normalize_list([0, 50, 100])
assert abs(r[0]-0.0) < 0.01 and abs(r[1]-0.5) < 0.01 and abs(r[2]-1.0) < 0.01
print("✓ Exercise 4.3 passed!")
```

Output:
```
✓ Exercise 4.3 passed!
```

### Exercise 4.4: Sigmoid Function


<!-- grade_id: ex4_4, points: 0, grade: false, solution: true -->
```python
import numpy as np
def simple_sigmoid(x):
    """
    Sigmoid: σ(x) = 1 / (1 + e^(-x))
    """
    # YOUR CODE HERE
    return 1 / (1 + np.exp(-x))


# Test
for x in [-5, 0, 5]:
    print(f"sigmoid({x}) = {simple_sigmoid(x):.4f}")
```

Output:
```
sigmoid(-5) = 0.0067
sigmoid(0) = 0.5000
sigmoid(5) = 0.9933
```

<!-- grade_id: ex4_4_test, points: 1, grade: true, solution: false -->
```python
assert abs(simple_sigmoid(0) - 0.5) < 0.01
assert simple_sigmoid(10) > 0.9 and simple_sigmoid(-10) < 0.1
print("✓ Exercise 4.4 passed!")
```

Output:
```
✓ Exercise 4.4 passed!
```

---

## Final Challenge: Complete ML Pipeline without a model

Combine everything you learned!

<div dir="rtl">

## التحدي النهائي: بناء ML Pipeline كامل بدون موديل

</div>

<div dir="rtl">

اجمع كل ما تعلمته!

</div>


<!-- grade_id: final, points: 0, grade: false, solution: true -->
```python
# Given data
data_points = [1.5, 4.2, 2.8, 5.6, 3.1]
true_labels = [0, 1, 0, 1, 1]

# Task 1: Make predictions (if > 3.0, predict 1, else 0)
def make_predictions(data_points):
  prediction =[]
  for i in range(len(data_points)):
    prediction.append(int(data_points[i] > 3.0))
  return prediction


# Task 2: Calculate accuracy
# YOUR CODE HERE
def accuracy(predictions, true_labels):
  correct = 0
  for i in range(len(predictions)):
    if predictions[i] == true_labels[i]:
      correct += 1
  accuracy = (correct / len(predictions)) * 100
  return accuracy

# Task 3: Normalize data
# YOUR CODE HERE
def normalized(data_points):
  normalized = []
  min_val = min(data_points)
  max_val = max(data_points)
  range_val = max_val - min_val
  for val in data_points:
    norm_val = (val - min_val) / range_val
    normalized.append(norm_val)
  return normalized

normalized_data = normalized(data_points)
predictions = make_predictions(data_points)
accuracy = accuracy(predictions, true_labels)

# Results
print("╔═══════════════════════════════╗")
print("║   ML Pipeline Results        ║")
print("╚═══════════════════════════════╝")
print(f"\nData: {data_points}")
print(f"Normalized: {[f'{x:.3f}' for x in normalized_data]}")
print(f"\nTrue: {true_labels}")
print(f"Pred: {predictions}")
print(f"\nAccuracy: {accuracy:.1f}%")
```

Output:
```
╔═══════════════════════════════╗
║   ML Pipeline Results        ║
╚═══════════════════════════════╝

Data: [1.5, 4.2, 2.8, 5.6, 3.1]
Normalized: ['0.000', '0.659', '0.317', '1.000', '0.390']

True: [0, 1, 0, 1, 1]
Pred: [0, 1, 0, 1, 1]

Accuracy: 100.0%
```

<!-- grade_id: final_test, points: 2, grade: true, solution: false -->
```python
assert predictions == [0, 1, 0, 1, 1] and accuracy == 100.0
assert len(normalized_data) == 5 and all(0 <= x <= 1 for x in normalized_data)
print("✓ Final challenge passed!")
print("\n🎉 You built a complete ML pipeline!")
```

Output:
```
✓ Final challenge passed!

🎉 You built a complete ML pipeline!
```